## Math Practice

<!-- grade_id: init, points: 0, grade: true, solution: false -->
```python
#########################################################
### EXECUTE THIS CELL BEFORE YOU  TEST YOUR SOLUTIONS ###
#########################################################

import imp
sol = imp.load_compiled("sol", "./sol.py")
from nose.tools import assert_equal
```

## A note on variable initialization

Variable initialization refers to the act of declaring a variable (giving it a name) and assigning it a value. A variable's name is declared on the left side of the `=` and its value on the right side. This can be done in a single line of code, like the below. Note that a variable's name may not contain spaces or special characters, with the exception of the underscore `_`.

```python
my_variable = 1
```

Once initialized, referring to `my_variable` in subsequent lines of code will actually refer to `my_variable`'s value, which has been set to `1`. However, as their name implies, variables can have changing values. For example, we might do something like the below:

```python
my_variable = 1
print(my_variable)     # This will print out the value of my_variable
# >>> 1

my_variable = 1 * 2    # We've now changed the value of my_variable to be (1 * 2) = 2
print(my_variable)
# >>> 2
```

We can also reference a variable's current value in it's own re-assignment! This can be done if the variable has already been assigned a value previously. For example:
```python
my_variable = my_variable * 3
print(my_variable)
# >>> 6
```
What the above line is doing is setting the value of `my_variable` equal to the current value of `my_variable`, which we set to `2`, multiplied by `3`. This ends up being: `my_variable = 2 * 3 = 6`

Calculate the result of 3.93 multiplied by 4901 and save it in a variable named 'q1', then print it out.

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
# your code here
q1 = 3.93 * 4901
print(q1)
```

Output:
```
19260.93
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
# Note: All test cases are read-only

assert_equal(q1, sol.q1)
print("Success!")
```

Output:
```
Success!
```

Calculate the result of 215 divided by 6 and save it in a variable named 'q2', then print it out.

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
# your code here
q2 = 215/6
print(q2)
```

Output:
```
35.833333333333336
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q2, sol.q2)
print("Success!")
```

Output:
```
Success!
```

Calculate the result of 3 divided by 0.3 and save it in a variable named 'q3', then print it out.

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
# your code here
q3 = 3/0.3
print(q3)
```

Output:
```
10.0
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q3, sol.q3)
print("Success!")
```

Output:
```
Success!
```

Calculate the remainder of 215 divided by 6 and save it in a variable named 'q4', then print it out.

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
# your code here
q4 = 215%6
print(q4)
```

Output:
```
5
```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q4, sol.q4)
print("Success!")
```

Output:
```
Success!
```

Calculate the value of 9 raised to the 12th power and save it in a variable named 'q5', then print it out.

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
# your code here
q5 = 9**12
print(q5)
```

Output:
```
282429536481
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q5, sol.q5)
print("Success!")
```

Output:
```
Success!
```

Cast 3.5 to an integer and save it in a variable named 'q6', then print it out.

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
# your code here
q6 = int(3.5)
print(q6)
```

Output:
```
3
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q6, sol.q6)
print("Success!")
```

Output:
```
Success!
```

## Data Types Practice

Calculate the data type of `"False"` (notice the **quotes** around the word False!) and save it in a variable named 'q7', then print it out.

<!-- grade_id: q7, points: 0, grade: false, solution: true -->
```python
# your code here
q7 = type("False")
print(q7)
```

Output:
```
<class 'str'>
```

<!-- grade_id: q7_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q7, sol.q7)
print("Success!")
```

Output:
```
Success!
```

Calculate the data type of True and save it in a variable named 'q8', then print it out.

<!-- grade_id: q8, points: 0, grade: false, solution: true -->
```python
# your code here
q8 = type(True)
print(q8)
```

Output:
```
<class 'bool'>
```

<!-- grade_id: q8_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q8, sol.q8)
print("Success!")
```

Output:
```
Success!
```

Calculate the data type of the result of 1000 divided by 10 and save it in a variable named 'q9', then print it out.

<!-- grade_id: q9, points: 0, grade: false, solution: true -->
```python
# your code here
q9 = type(1000/10)
print(q9)
```

Output:
```
<class 'float'>
```

<!-- grade_id: q9_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q9, sol.q9)
print("Success!")
```

Output:
```
Success!
```

Cast the value of 6.3 divided by 3.8 to an integer. Save it in a variable named 'q10', then print it out.

<!-- grade_id: q10, points: 0, grade: false, solution: true -->
```python
# your code here
q10 = int(6.3/3.8)
print(q10)
```

Output:
```
1
```

<!-- grade_id: q10_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q10, sol.q10)
print("Success!")
```

Output:
```
Success!
```

## String Practice

Concatenate the strings 'James', 'Brian', and 'Patrick' - store the result in a variable called 'q11', then print it out. Make sure to add a single empty space between the names!

<!-- grade_id: q11, points: 0, grade: false, solution: true -->
```python
# your code here
q11 = 'James'+ ' ' + 'Brian'+' ' +'Patrick'
print(q11)
```

Output:
```
James Brian Patrick
```

<!-- grade_id: q11_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q11, sol.q11)
print("Success!")
```

Output:
```
Success!
```

Make the following string correct, storing it in a variable named 'q12':<br />
`q12 = "4 % 2 = " + (4 % 2))`

<!-- grade_id: q12, points: 0, grade: false, solution: true -->
```python
# your code here
q12 = str(4 % 2)
print(q12)
```

Output:
```
0
```

<!-- grade_id: q12_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q12, sol.q12)
print("Success!")
```

Output:
```
/opt/conda/lib/python3.7/unittest/case.py in assertMultiLineEqual(self, first, second, msg)
   1231             diff = '\n' + ''.join(difflib.ndiff(firstlines, secondlines))
   1232             standardMsg = self._truncateMessage(standardMsg, diff)
-> 1233             self.fail(self._formatMessage(msg, standardMsg))
   1234 
   1235     def assertLess(self, a, b, msg=None):

/opt/conda/lib/python3.7/unittest/case.py in fail(self, msg)
    691     def fail(self, msg=None):
    692         """Fail immediately, with the given message."""
--> 693         raise self.failureException(msg)
    694 
    695     def assertFalse(self, expr, msg=None):

AssertionError: '0' != '4 % 2 = 0'
- 0
+ 4 % 2 = 0
```

Save the following quote, (including the double quotes), in a variable called 'q13', then print it:<br />
Albert Einstein's best quote is "I have no special talent. I am only passionately curious."

<!-- grade_id: q13, points: 0, grade: false, solution: true -->
```python
# your code here
q13 = '"I have no special talent. I am only passionately curious."'
print(q13)
```

Output:
```
"I have no special talent. I am only passionately curious."
```

<!-- grade_id: q13_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(q13, sol.q13)
print("Success!")
```

Output:
```
/opt/conda/lib/python3.7/unittest/case.py in assertMultiLineEqual(self, first, second, msg)
   1231             diff = '\n' + ''.join(difflib.ndiff(firstlines, secondlines))
   1232             standardMsg = self._truncateMessage(standardMsg, diff)
-> 1233             self.fail(self._formatMessage(msg, standardMsg))
   1234 
   1235     def assertLess(self, a, b, msg=None):

/opt/conda/lib/python3.7/unittest/case.py in fail(self, msg)
    691     def fail(self, msg=None):
    692         """Fail immediately, with the given message."""
--> 693         raise self.failureException(msg)
    694 
    695     def assertFalse(self, expr, msg=None):

AssertionError: '"I have no special talent. I am only pass[14 chars]us."' != 'Albert Einstein\'s best quote is "I have [47 chars]us."'
- "I have no special talent. I am only passionately curious."
+ Albert Einstein's best quote is "I have no special talent. I am only passionately curious."
? ++++++++++++++++++++++++++++++++
```

```python

```